rule_files:
- ../aggregate-prometheus-rules/fluent-bit.rules.yaml

evaluation_interval: 30s

tests:

- interval: 30s
  external_labels:
    seed: aws
  input_series:
  # FluentBitDown
  - series: 'up{job="fluent-bit"}'
    values: '0+0x30'
  alert_rule_test:
  - eval_time: 15m
    alertname: FluentBitDown
    exp_alerts:
    - exp_labels:
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: >
          There are no fluent-bit pods running on seed: aws.
          No logs will be collected.
        summary: Fluent-bit is down

- interval: 1m
  external_labels:
    seed: aws
  input_series:
  # FluentBitDown
  - series: 'fluentbit_input_bytes_total{pod="fluent-bit-test"}'
    values: '1+1x3 4+0x370'
  alert_rule_test:
  - eval_time: 370m
    alertname: FluentBitIdleInputPlugins
    exp_alerts:
    - exp_labels:
        service: logging
        severity: warning
        type: seed
        visibility: operator
        pod: fluent-bit-test
      exp_annotations:
        description: >
          The input plugins of Fluent-bit pod fluent-bit-test running on seed aws haven't collected any logs for the last 6 hours.
        summary: Fluent-bit input plugins haven't process any data for the past 6 hours

- interval: 1m
  external_labels:
    seed: aws
  input_series:
  # FluentBitReceivesLogsWithoutMetadata
  - series: 'fluentbit_loki_gardener_logs_without_metadata_total{pod="fluent-bit-test"}'
    values: '0+0x3 0+1x30'
  alert_rule_test:
  - eval_time: 22m
    alertname: FluentBitReceivesLogsWithoutMetadata
    exp_alerts:
    - exp_labels:
        pod: fluent-bit-test
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: >
          fluent-bit-test receives logs without metadata on seed:
          aws. These logs will be dropped.
        summary: Fluent-bit receives logs without metadata

- interval: 1m
  external_labels:
    seed: aws
  input_series:
  # FluentBitSendsOoOLogs
  - series: 'prometheus_target_scrapes_sample_out_of_order_total{pod="fluent-bit-test"}'
    values: '0+0x3 0+1x30'
  alert_rule_test:
  - eval_time: 22m
    alertname: FluentBitSendsOoOLogs
    exp_alerts:
    - exp_labels:
        pod: fluent-bit-test
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: >
          fluent-bit-test on seed: aws sends OutOfOrder logs
          to the Loki. These logs will be dropped.
        summary: Fluent-bit sends OoO logs

- interval: 1m
  external_labels:
    seed: aws
  input_series:
  # FluentBitGardenerLokiPluginErrors
  - series: 'fluentbit_loki_gardener_errors_total{pod="fluent-bit-test"}'
    values: '0+0x3 0+1x30'
  alert_rule_test:
  - eval_time: 22m
    alertname: FluentBitGardenerLokiPluginErrors
    exp_alerts:
    - exp_labels:
        pod: fluent-bit-test
        service: logging
        severity: warning
        type: seed
        visibility: operator
      exp_annotations:
        description: >
          There are errors in the fluent-bit-test GardenerLoki plugin on seed:
          aws.
        summary: Errors in Fluent-bit GardenerLoki plugin

